# Assignment 1 
### Neural Networks and Fuzzy Logic (2nd Semester 2019-2020)

### Instructions:
This assignment is designed to check your understanding of linear algebra, assuming you have basic knowlegde of Linear Algebra, this assignment will introudce you to some high level tasks in Linear Algebra, and what's the fun part about it? You'll be implemnting those concepts and visualizing them!! So Let's start.

#### Key Concepts: Matrix Inverse, Linear Transforms, Eigendecomposition and all the interesting stuff Linear Algebra has :)

##### Please make sure you adhere to the policy of originality of your work, the assignment should reflect your understanding and should be your original, cases of plagiarism will be strictly dealt with as per institute norms. The penalties can include debarrment from the course and registration for the next semester.

# Matrices as linear transformations

Matrices can be viewed as linear transformations you 'apply' to vectors. Some matrices will rotate your space, others will rescale it etc. So when we apply a matrix to a vector, we end up with a transformed version of the vector. When we say that we 'apply' the matrix to the vector it means that we calculate the dot product of the matrix with the vector. We will start with a basic example of this kind of transformation.

```python
#No other imports are allowed !!!
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
```

```python
#utility function, please don't alter this
def plot_vectors(vectors, cols, alpha=1):
    """
    Plot set of vectors.
    Parameters
    ----------
    vectors : array-like
        Coordinates of the vectors to plot. Each vectors is in an array. For
        instance: [[1, 3], [2, 2]] can be used to plot 2 vectors.
    cols : array-like
        Colors of the vectors. For instance: ['green', 'red'] will display the
        first vector in green and the second in red.
    alpha : float
        Opacity of vectors
    Returns:
    fig : instance of matplotlib.figure.Figure
        The figure of the vectors
    """
    plt.axvline(x=0, color='#A9A9A9', zorder=0)
    plt.axhline(y=0, color='#A9A9A9', zorder=0)

    for i in range(len(vectors)):
        if (isinstance(alpha, list)):
            alpha_i = alpha[i]
        else:
            alpha_i = alpha
        x = np.concatenate([[0,0],vectors[i]])
        plt.quiver([x[0]],
                   [x[1]],
                   [x[2]],
                   [x[3]],
                   angles='xy', scale_units='xy', scale=1, color=cols[i],
                  alpha=alpha_i)
```

```python
A = np.array([[1, 3], [2, 6]])
print(A)
print(A.shape)
```

```python
v = np.array([[5], [3]])
print(v)
print(v.shape)
```

```python
plot_vectors([v.flatten()],['red'])
plt.ylim(0, 6)
plt.xlim(0, 6)
plt.show()
```

**Inline Question**: did you see that v is flattened before plotting? why do we need to do that? Check out the code for plot_vectors and try to answer this? <br>

<!-- grade_id: cell-89c617592276dccd, points: 1, grade: true, solution: true -->
Answer:

<!-- grade_id: cell-8c70805160037e04, points: 0, grade: false, solution: true -->
```python
def apply_linear_transform(v,A):
    '''
    Apply the linear transformation on v by a matrix A
    return the transformed matrix (Av)
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: cell-4f7d4f29ea739491, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR apply_linear_transform
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
mat0 = np.random.randn(10)
mat1 = np.random.randn(10,10)
mat2 = np.random.randn(10,10,10)

vect1 = np.random.randn(10,10,10)
vect2 = np.random.randn(10,10)
vect3 = np.random.rand(10)

assert np.array_equal(apply_linear_transform(mat0,vect3), mat0.dot(vect3))
assert np.array_equal(apply_linear_transform(mat1,vect2), mat1.dot(vect2))
assert np.array_equal(apply_linear_transform(mat2,vect1), mat2.dot(vect1))
### END HIDDEN TESTS
```

```python

# Let us visualize the transformation on vector v by a matrix A as mentioned above
transformed_vector = apply_linear_transform(v,A)
plot_vectors([v.flatten(), transformed_vector.flatten()],cols=['red','green'])
plt.ylim(0, 6)
plt.xlim(0, 6)
plt.show()
```

## Eigendecomposition

We have seen an example of a vector transformed by a matrix. Now imagine that the transformation of the initial vector gives us a new vector that has the exact same direction. The scale can be different but the direction is the same. Applying the matrix didn't change the direction of the vector. This special vector is called an eigenvector of the matrix. We will see that finding the eigenvectors of a matrix can be very useful.

<span class='pquote'>
    Imagine that the transformation of the initial vector by the matrix gives a new vector with the exact same direction. This vector is called an eigenvector of ${A}$.
</span>

This means that ${v}$ is a eigenvector of ${A}$ if ${v}$ and ${Av}$ are in the same direction or to rephrase it if the vectors ${Av}$ and ${v}$ are parallel. The output vector is just a scaled version of the input vector. This scalling factor is $\lambda$ which is called the **eigenvalue** of ${A}$.

$$
{Av} = \lambda{v}
$$

Talking in terms of eigen values and eigen vector the eigen decomposition of a matrix looks like the following:

$$
{A}= {V}\cdot{diag}({\lambda})\cdot{V}^{-1}
$$

We will also verify this property of eigendecomposition.

Eigen decomposition is a crucial step for PCA (introduced later)

Let's begin!!

<!-- grade_id: cell-947cfea731ce6537, points: 0, grade: false, solution: true -->
```python
def find_eigen_decomposition(matrix):
    '''
    use numpy to find the eigen decomposition of the matrix
    returns eigenValues and the eigenVectors
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: cell-9e1377220b4336dd, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR find_eigen_decomposition

### BEGIN HIDDEN TESTS
matrix = np.random.randn(10,10)
matrix1 = np.random.randn(100,10)
matrix2 = np.random.randn(20,10)
eigenValues,eigenVectors = np.linalg.eig(matrix)
eigenValues1,eigenVectors1 = np.linalg.eig(matrix1)
eigenValues2,eigenVectors2 = np.linalg.eig(matrix2)

from_func_eig,from_func_eig11 = find_eigen_decomposition(matrix)
from_func_eig1,from_func_eig111 = find_eigen_decomposition(matrix1)
from_func_eig2,from_func_eig222 = find_eigen_decomposition(matrix2)

assert np.array_equal(eigenValues,from_func_eig)
assert np.array_equal(eigenValues1,from_func_eig1)
assert np.array_equal(from_func_eig2,eigenValues2)
### END HIDDEN TESTS
```

<!-- grade_id: cell-b7f2c7683aef8586, points: 0, grade: false, solution: true -->
```python
def verify_eigen_decomposition(matrix,eigen_values,eigen_vectors):
    
    '''
    verifies the eigen decomposition property of matrix
    return the recontsructed matrix using the property of eigen vectors and eigenvalues
    '''
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

<!-- grade_id: cell-ca924fb2cbe495d7, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR verify_eigen_decomposition

### BEGIN HIDDEN TESTS
matrix = np.random.randn(10,10)
eigenValues,eigenVectors = np.linalg.eig(matrix)
assert np.array_equal(verify_eigen_decomposition(matrix,eigenValues,eigenVectors),matrix)
### END HIDDEN TESTS
```

## Principal Component Analysis
Handling Dimensions is very important in Machine Leanring. The dimensions are all the features of the dataset. For instance, if you are looking at a dataset containing pieces of music, dimensions could be the genre, the length of the piece, the number of instruments, the presence of a singer etc. You can imagine all these dimensions as different columns. When there is only two dimensions, it is very convenient to plot: you can use the $x$ - and  $y$ -axis. Add color and you can represent a third dimension. It is similar if you have tens or hundereds of dimensions, it will just be harder to visualize it.

When you have that many dimensions it happens that some of them are correlated. For instance, we can reasonably think that the genre dimension will correlate with the instruments dimensions in our previous example. One way to reduce dimensionality is simply to keep only some of them. The problem is that you loose good information. It would be nice to have a way to reduce these dimensions while keeping all the information present in the data set.

The aim of principal components analysis (PCA) is generaly to reduce the number of dimensions of a dataset where dimensions are not completly decorelated. PCA provides us with a new set of dimensions, the principal components (PC).

Step-wise PCA is:

1) Center the data <br>
2) Find the covariance matrix <br>
3) Find the Eigen Decompostion <br>
4) Choose the principal components <br>
5) Find the new dataset Matrix <br>

Following [article](https://towardsdatascience.com/a-one-stop-shop-for-principal-component-analysis-5582fb7e0a9c) also explains PCA very well.
Let's implement PCA to get an essence of how dimensions can be reduced, you need to implement the function one by one!

```python

x = 5*np.random.rand(100)
y = 2*x + 1 + np.random.randn(100)

#the two dimensions of our dataset
x = x.reshape(100, 1)
y = y.reshape(100, 1)

#The dataset matrix
X = np.hstack([x, y])
```

```python
plt.plot(X[:,0],X[:,1],'.')
plt.show()
```

<!-- grade_id: cell-4aace02eb82ebd5c, points: 0, grade: false, solution: true -->
```python
def center(X):
    '''
    function to center the data around the mean
    modify the return value to return the modified matrix
    '''
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

<!-- grade_id: cell-9e373b3dc8185259, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR center

### BEGIN HIDDEN TESTS

mat1 = np.random.randn(10,10)
mat2 = np.random.randn(100,100)

mat11 = mat1 - np.mean(mat1,axis=0)
mat22 = mat2 - np.mean(mat2,axis=0)

assert np.array_equal(mat11,center(mat1))
assert np.array_eqaul(mat22,center(mat2))

### END HIDDEN TESTS
```

<!-- grade_id: cell-f2ba4844ba63791f, points: 0, grade: false, solution: true -->
```python
def find_covariance_matrix(X):
    '''
    Find the covariance matrix of X
    Return the final matrix
    '''
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

<!-- grade_id: cell-23455edbef9a7b84, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR find_covariance_matrix

### BEGIN HIDDEN TESTS
mat1 = np.random.randn(10,10)
mat2 = np.random.randn(100,100)

assert np.array_equal(np.cov(mat1),find_covariance_matrix(mat1))
assert np.array_equal(np.cov(mat2),find_covariance_matrix(mat2))
### END HIDDEN TESTS
```

Now let's plot the centered data and the eigen vectors in the same figure

```python
# We will use the simulated X, the one we generated in one of the above cells
X_centered = center(X)
covariance = find_covariance_matrix(X_centered)
eigenVals, eigenVecs = find_eigen_decomposition(covariance)


plot_vectors(eigenVecs.T, ['red', 'blue'])
plt.plot(X_centered[:,0], X_centered[:,1], '*')
plt.xlim(-3, 3)
plt.ylim(-3, 3)
plt.show()
```

<!-- grade_id: cell-365c5fe2dd253c56, points: 0, grade: false, solution: true -->
```python
def two_dimensional_pca(X_centered, eigenValues, eigenVecs):
    '''
    using X_centered, eigenValues, eigenVecs as obtained above
    apply pca on X_centered and chose the principal component
    return the new matrix
    '''
    
    ### BEGIN SOLUTION
    
    ### END SOLUTION
```

<!-- grade_id: cell-f8a28e82ea7c8e1d, points: 1, grade: true, solution: false -->
```python
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR two_dimensional_pca

### BEGIN HIDDEN TESTS
X_new = two_dimensional_pca(X_centered, eigenVals, eigenVecs)
assert np.array_equal(X_new,X)
### END HIDDEN TESTS
```

**Inline Questions:** <br> 
1) Can you identify the color of the line which points to the direction of maximum variance? <br>
2) Look at the Graph below and explain if PCA can be applied on this dataset or not? Do give reason. <br>
3) Can you apply PCA to categorical variables? (Answer to this will give you hints for the next excercise)

![Plot](images/pca_inline.png)

Answers:

## PCA on a large dataset
## About Dataset:
Task: Reducing the dimesions for predicting house prices.
The dataset is stored in the file **'data_pca.csv'**, it has 15 numerical columns and 2 categorical columns (columns are labeled as $X0$....$X16$), in an imaginary world, all these columns have certain significance towards prediciton of final house prices. Your task is simple, apply PCA to reduce the number of dimensions. 

Now implement the following functions for getting the final dataset with reduced dimensions. Keep the top $k (k=0...10)$ principal components.

<!-- grade_id: cell-b2744694d03e082f, points: 0, grade: false, solution: true -->
```python
def read_data(name):
    '''
    read the CSV file and return the corresponding dataframe
    don't forget to remove the index column (check the name properly)
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: cell-7c4fe4ab88e97f99, points: 1, grade: true, solution: false -->
```python
#input the path top the CSV file
#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR read_data

PATH_TO_CSV_FILE = './data_pca.csv'
### BEGIN HIDDEN TESTS
df = pd.read_csv(PATH_TO_CSV_FILE).drop(['Unnamed: 0'],axis=1)
df2 = read_data(PATH_TO_CSV_FILE)
assert np.array_equal(df2.to_numpy().flatten(),df.to_numpy().flatten())
### END HIDDEN TESTS
```

<!-- grade_id: cell-b944aaa1481e4ea8, points: 0, grade: false, solution: true -->
```python
def apply_pca(data,k):
    '''
    data: a pandas data frame
    Implement PCA on the dataset and return the transformed dataset as a pandas dataframe
    keep k pricipal components
    remember the dataset MUST have the chosen components and the columns you decide to exclude 
    from PCA calculation.
    The returned dataframe should NOT have the index column
    The returned dataframe should have similar column names as original and must be in the 
    SAME ORDER as in the original data, example : if X0 and X1 are chosen then X0 must appear
    appear before X1 in the returned frame as it appears before in the original frame as well.
    
    Try to verify the function on your own :)
    '''
    ### BEGIN SOLUTION
    ### END SOLUTION
```

<!-- grade_id: cell-4db7e2b7485907b8, points: 4, grade: true, solution: false -->
```python
PATH_TO_CSV_FILE = './data_pca.csv'
assert np.array_equal(pd.read_csv('pca_4.csv').to_numpy(),apply_pca(data,4).to_numpy())

#DO NOT ALTER THIS CELL
#HIDDEN TEST CASES FOR apply_pca
### BEGIN HIDDEN TESTS
assert np.array_equal(pd.read_csv('pca_5.csv').to_numpy(),apply_pca(data,5).to_numpy())
assert np.array_equal(pd.read_csv('pca_7.csv').to_numpy(),apply_pca(data,7).to_numpy())
assert np.array_equal(pd.read_csv('pca_10.csv').to_numpy(),apply_pca(data,10).to_numpy())
### END HIDDEN TESTS
```